# Aula 01 - Python Não Tão Básico

# Exercício 01:

Crie uma função que, dado duas listas sendo uma com nomes e outras com números, retorne uma __lista__ com os nomes associados aos números ordenados em ordem decrescente. Caso mais de um nome esteja associado a um mesmo número, eles devem ser ordenados em ordem lexicográfica. Por exemplo, dado uma lista com o nome de alguns jogadores e outra lista com os salários desses jogadores, a lista com o nome dos jogadores deve ser ordenada de forma que os primeiros jogadores sejam aqueles que possuem os maiores salários. No exemplo dado, a saída esperada é: ['Messi', 'Cristiano Ronaldo', 'Neymar', 'Bale', 'Piqué', 'Kroos', 'Sergio Ramos', 'Ibrahimovic', 'Griezmann', 'Suarez', 'Iniesta', 'Lavezzi', 'Sanchez', 'Benzema', 'Muller', 'Cavani', 'Lewandowski', 'Oscar', 'Pogba', 'Neuer']


__Dicas:__  
não se esqueça de manter a relação existente entre o nome e o número, ou seja, o nome que ocupa a primeira posição da lista está associado ao número da primeira posição da outra lista.


<!-- grade_id: cell1-a00, points: 0, grade: false, solution: true -->
```python
player = ["Neymar", "Piqué", "Benzema", "Neuer", "Iniesta", "Sergio Ramos", "Messi", "Oscar", "Griezmann", "Cavani", "Bale", "Pogba", "Cristiano Ronaldo", "Sanchez", "Kroos", "Lewandowski", "Lavezzi", "Ibrahimovic", "Muller", "Suarez"]
player_salary = [81.5, 29, 23.5, 20.8, 25.5, 27.5, 126, 22.2, 26, 22.5, 44, 22, 96, 23.55, 28.2, 22.2, 25.3, 26.8, 23, 26]
from operator import itemgetter, attrgetter


def sort_lists(list1, list2):
    
    list3 = list(zip(list1, list2))
    wc1 = sorted(list3,
                    key=lambda wcs1: wcs1[0],
                    reverse=False)
    wc2 = sorted(wc1,
                    key=lambda wcs2: wcs2[1],
                    reverse=True)
    
    player_new, salary_new = zip(*wc2)
    return list(player_new)
   
```

<!-- grade_id: cell1t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
"""Check that sort_lists returns the correct output for several inputs"""

player = ["Neymar", "Piqué", "Benzema", "Neuer", "Iniesta", "Sergio Ramos", "Messi", "Oscar", "Griezmann", "Cavani", "Bale", "Pogba", "Cristiano Ronaldo", "Sanchez", "Kroos", "Lewandowski", "Lavezzi", "Ibrahimovic", "Muller", "Suarez"]
player_salary = [81.5, 29, 23.5, 20.8, 25.5, 27.5, 126, 22.2, 26, 22.5, 44, 22, 96, 23.55, 28.2, 22.2, 25.3, 26.8, 23, 26]

assert_equal(sort_lists(player, player_salary), ['Messi', 'Cristiano Ronaldo', 'Neymar', 'Bale', 'Piqué', 'Kroos', 'Sergio Ramos', 'Ibrahimovic', 'Griezmann', 'Suarez', 'Iniesta', 'Lavezzi', 'Sanchez', 'Benzema', 'Muller', 'Cavani', 'Lewandowski', 'Oscar', 'Pogba', 'Neuer'])
```

<!-- grade_id: cell1t2-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:

Crie uma função que remova todas as vogais sem acentuação de uma frase dada e retorne uma __lista__ (que pode ser vazia) em ordem lexicográfica com as palavras (sem vogais) da frase que possuem vogais acentuadas não removidas. Para testar sua função, você pode utilizar a entrada abaixo, cuja saída esperada é: ['dsprcrá', 'já', 'lgém', 'nxplcávl.', 'sbsttíd', 'srá', 'stá'].

Frase: "Existe uma teoria que diz que, se um dia alguém descobrir exatamente para que serve o Universo e por que ele está aqui, ele desaparecerá instantaneamente e será substituído por algo ainda mais estranho e inexplicável. Existe uma segunda teoria que diz que isso já aconteceu."


__Dicas:__  
não se preocupe em remover a pontuação do texto.  
a função unidecode(frase) tranforma todos os caracteres da frase em caracteres não acentuados. (Atenção: para usá-la import o módulo da seguinte forma: _from unidecode import unidecode_)

<!-- grade_id: cell2-import, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell2-a00, points: 0, grade: false, solution: true -->
```python
#Exemplo:
sample_text = sample_text = "Existe uma teoria que diz que, se um dia alguém descobrir exatamente para que \
serve o Universo e por que ele está aqui, ele desaparecerá instantaneamente e será substituído por algo ainda \
mais estranho e inexplicável. Existe uma segunda teoria que diz que isso já aconteceu."

import unidecode

def text_cons(sentence):
    unidecode(sentence)
    # YOUR CODE HERE
    #raise NotImplementedError()

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-16-10f2573a54a3> in <module>()
      2 sample_text = sample_text = "Existe uma teoria que diz que, se um dia alguém descobrir exatamente para que serve o Universo e por que ele está aqui, ele desaparecerá instantaneamente e será substituído por algo ainda mais estranho e inexplicável. Existe uma segunda teoria que diz que isso já aconteceu."
      3 
----> 4 import unidecode
      5 
      6 def text_cons(sentence):

ModuleNotFoundError: No module named 'unidecode'
```

<!-- grade_id: cell2t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
"""Check that text_cons returns the correct output for several inputs"""

sample_text = "Existe uma teoria que diz que, se um dia alguém descobrir exatamente para que \
serve o Universo e por que ele está aqui, ele desaparecerá instantaneamente e será substituído por algo ainda \
mais estranho e inexplicável. Existe uma segunda teoria que diz que isso já aconteceu."

assert_equal(text_cons(sample_text), ['dsprcrá', 'já', 'lgém', 'nxplcávl.', 'sbsttíd', 'srá', 'stá'])
```

<!-- grade_id: cell2t2-a00, points: 1, grade: true, solution: false -->
```python

```