# Aula 04 - Estatística

# Exercício 01:


Utilizando os dados de acidentes fatais em companhias aéreas dos Estados Unidos de 1985 a 1999, calcule algumas estatísticas básicas. Você deve retornar uma __lista__ com os valores das estatísticas calculadas, sendo elas, nessa ordem: menor valor, maior valor, média, mediana, dispersão, variância, desvio-padrão e intervalo-interquartil. Para responder à questão, crie uma função _estat(acidentes)_ que retorna a lista com os valores correspondentes às resposta da questão, inteiros quando forem inteiros ou com uma precisão de 4 casas decimais caso contrário.

__Teste:__
`assert_equal(estat(acidentes), ans)`, sendo que `ans` é uma lista contendo os valores corretos para as estatísticas que este exercício pede.

__Dicas:__  
1) A função round(var, n) retorna o valor da variável var com uma precisão de n casas decimais.

2) Execute o teste `assert_equal(estat(lista_boba), ans_bobo)` para alguma `lista_boba` que você saiba calcular as estatísticas no papel.


__Fonte:__ https://aviation-safety.net/

```python
import numpy as np
import math
```

```python
 def mean(x):
    return sum(x) / len(x)
```

```python
def median(v):
    n = len(v)
    sorted_v = sorted(v)
    return (sorted_v[math.floor((n-1)/2)] + sorted_v[math.ceil((n-1)/2)])/2
```

```python
def data_range(x):
    return max(x) - min(x)
```

```python
def de_mean(x):
    x_bar = mean(x)
    return [x_i - x_bar for x_i in x]

def variance(x):
    if(len(x) > 2):
        n = len(x)
        deviations = de_mean(x)
        deviations = np.array(deviations) 
        return np.sum(deviations ** 2) / (n-1)
    return None
```

```python
def standard_deviation(x):
    return math.sqrt(variance(x))
```

```python
def quantile(x, p):
    p_index = int(p * len(x))
    return sorted(x)[p_index]

def interquartile_range(x):
    return quantile(x, 0.75) - quantile(x, 0.25)
```

<!-- grade_id: cell1-a00, points: 0, grade: false, solution: true -->
```python
# Crie aqui a função estat(acidentes) - com esse nome e parâmetro - 
# a função deve retornar a lista com as respostas com precisão de 4 casas decimais!

acidentes = [0,14,0,1,0,4,1,0,0,2,1,5,0,3,0,0,6,1,1,12,3,1,5,0,3,0,0,1,1,0,1,5,2,1,1,3,4,0,3,0,2,2,1,0,1,1,1,3,0,4,3,8,7,3,0,1]

list_boba = [0,1,2,3,4,5]

def estat(data_set):
    
    estats_data_set =  np.ndarray((8,1),dtype=float)
    
    estats_data_set[0] = min(data_set)
    estats_data_set[1] = max(data_set)
    estats_data_set[2] = round(mean(data_set),4)
    estats_data_set[3] = round(median(data_set),4)
    estats_data_set[4] = round(data_range(data_set),4)
    estats_data_set[5] = round(variance(data_set),4)
    estats_data_set[6] = round(standard_deviation(data_set),4)
    estats_data_set[7] = round(interquartile_range(data_set),4)
    
    return estats_data_set

```

<!-- grade_id: cell1t-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:


Procure encontrar correlações espúrias interessantes e apresente um exemplo encontrado. Ou seja, aprensente dois conjuntos de dados que possuem alta correlação (muito positivas ou muito negativas) sem que um seja de fato o causador do outro. Além disso, deixe resgistrado os gráficos com a distribuição dos dados e um gráfico de dispersão como forma de visualizar a correlação entre os dados. Calcule a covariância e correlação entre os dados e, por fim, se possível, tente explicar qual poderia ser a verdadeira causa da ocorrência das observações. Para isso, utilize a última célula desse notebook.


__Observação:__  
Para ideias de correlações espúrias, veja os seguintes sites:  
http://tylervigen.com/spurious-correlations  
https://en.wikipedia.org/wiki/Spurious_relationship#Other_relationships  


<!-- grade_id: cell21-a00, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell22-a00, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Exercício 03:

O Paradoxo de Simpson também pode ser observado na taxa de sobrevivência do [naufrágio do RMS Titanic](https://en.wikipedia.org/wiki/RMS_Titanic). Havia uma estimativa de 2.224 passageiros e tripulantes a bordo, e mais de 1.500 morreram, tornando-se um dos mais mortais desastres marítimos comerciais em tempos de paz da história moderna. Dados sobre os passageiros podem ser baixados [aqui](https://ww2.amstat.org/publications/jse/v3n3/datasets.dawson.html). Este conjunto de dados foi processado por mim e transformado em um arquivo [csv](https://www.dropbox.com/s/vk8jf0wyczqxkvv/survival_titanic.csv?dl=0), que é muito mais fácil de tratar que um arquivo texto. Informações sobre ele podem ser lidas [aqui](https://www.dropbox.com/s/xpjw74khyqx9ww4/survival_titanic.README.txt?dl=0). 

Embora muitas outras informações existam sobre os passageiros, aqui vamos trabalhar com apenas quatro:

1) Se o passageiro sobreviveu;  
2) A classe do seu bilhete (primeira, segunda ou terceira) ou se ele era membro da tripulação;  
3) O sexo do passageiro;  
4) E se ele era um adulto ou uma criança.  

Os dados estão organizados da seguinte maneira:

|Coluna|Descrição|Valores|
| :-: | :-: | :-: |
|0|Classe| 0 = tripulação, 1 = primeira, 2 = segunda, 3 = terceira|
|1|Idade|1 = adulto, 0 = criança|
|2|Sexo|1 = masculino, 0 = feminino|
|3|Sobreviveu?|1 = sim, 0 = não|


O seguinte trecho de código carrega os dados que serão usados para responder as questões a seguir:


```python
import pandas as pd
df = pd.read_csv('survival_titanic.csv')
X = df.values
```

__a)__ Verifique qual classe possui uma taxa de sobrevivência maior, para isso, implemente uma função _survivalRatePerClass(X, classe)_ com esse mesmo nome e parêmetros que, dado a tabela X e um número (0,1,2,3) representando a classe, retorna a proporção de pessoas daquela classe que sobreviveram com uma precisão de 5 casas decimais. 

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell31-a00, points: 0, grade: false, solution: true -->
```python
# from collections import Counter
import numpy as np
def survivalRatePerClass(X, classe):

    total_per_class = Counter(list(X[:,0]))
    qtd_lives = 0
    
    for row in X:
        if row[0] == classe and row[3] == 1:
            qtd_lives = qtd_lives + 1

    return round((qtd_lives/total_per_class[classe]),5)

survival_rate = { 0 : survivalRatePerClass(X, 0) ,
                  1 : survivalRatePerClass(X, 1) , 
                  2 : survivalRatePerClass(X, 2), 
                  3 :survivalRatePerClass(X, 3) }
print( "The class with the major survival rate is the class " + str(max(survival_rate, key=survival_rate.get) ) )


```

Output:
```
The class with the major survival rate is the class 1
```

<!-- grade_id: cell3t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(survivalRatePerClass(X, 0), 0.23955)
```

<!-- grade_id: cell3t2-a00, points: 1, grade: true, solution: false -->
```python

```

__b)__ Verifique qual classe possui uma taxa de sobrevivência maior de acordo com a idade e sexo dos passageiros. Para isso, implemente uma função _survivalRatePerClassAndColumn(X, classe, coluna)_ com esse mesmo nome e parêmetros que, dado a tabela X, um número (0,1,2,3) representando a classe e um número (1,2) representando respectivamente as colunas que representam a idade e o sexo dos passageiros, retorna a proporção de pessoas daquela classe que sobreviveram com uma precisão de 5 casas decimais. Nesse caso, retorne uma tupla do tipo (<proporção de sobreviventes crianças da classe>, <proporção de sobreviventes adultos da classe>) quando a coluna for a 1 e (<proporção de sobreviventes mulheres da classe>, <proporção de sobreviventes homens da classe>)  quando a coluna for a 2.

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell32-a00, points: 0, grade: false, solution: true -->
```python
def survivalRatePerClassAndColumn(X, classe, column):

    total_column_true = 0
    qtd_lives_column_true = 0
    
    total_column_false = 0
    qtd_lives_column_false = 0
    
    for row in X:
            #eh da classe  e da na coluna
        if row[0] == classe and row[column] == 1 :
            
            total_column_true += 1
            
            if  row[3] == 1 :
                qtd_lives_column_true += 1
            
        elif row[0] == classe and row[column] == 0:
            total_column_false += 1
            if row[3] == 1 :
                qtd_lives_column_false += 1
    
    return (
                None if total_column_false == 0  or  qtd_lives_column_false == 0   else round((qtd_lives_column_false/total_column_false),5),
                None if total_column_true == 0 or  qtd_lives_column_true == 0 else round((qtd_lives_column_true/total_column_true),5) 
            )


```

<!-- grade_id: cell3t3-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

#classe 0 (criança, adulto)
assert_equal(survivalRatePerClassAndColumn(X, 0, 1), (None, 0.23955))
#classe 1 (criança, adulto)
assert_equal(survivalRatePerClassAndColumn(X, 1, 1), (1.0, 0.61755))
#classe 0 (mulher, homem)
assert_equal(survivalRatePerClassAndColumn(X, 0, 2), (0.86957, 0.22274))
```

<!-- grade_id: cell3t4-a00, points: 1, grade: true, solution: false -->
```python

```

__c)__ Após brincar com as funções criadas, faça uma breve análise sobre os resultados encontrados, relacionando-os ao Paradoxo de Simpson visto em sala de aula.

```python
def totalPersonsPerClassAndColumn(X, classe, column):

    total_column_true = 0
    total_column_false = 0
    
    for row in X:
        if row[0] == classe and row[column] == 1 :
            total_column_true += 1
            
        elif row[0] == classe and row[column] == 0:
            total_column_false += 1
    
    return (
                None if total_column_false == 0  else total_column_false,
                None if total_column_true == 0  else total_column_true 
            )


print("Tripulantes :",survivalRatePerClass(X, 0) )
print("Classe 1 :",survivalRatePerClass(X, 1) )
print("Classe 2 :",survivalRatePerClass(X, 2) )
print("Classe 3 :",survivalRatePerClass(X, 3) )
print("")
print("Tripulantes :criancas x adultos",survivalRatePerClassAndColumn(X, 0, 1) )
print("Classe 1 :criancas x adultos",survivalRatePerClassAndColumn(X, 1, 1) )
print("Classe 2 :criancas x adultos",survivalRatePerClassAndColumn(X, 2, 1) )
print("Classe 3 :criancas x adultos",survivalRatePerClassAndColumn(X, 3, 1) )
print("")
print("Tripulantes : mulheres x homens",survivalRatePerClassAndColumn(X, 0, 2) )
print("Classe 1 : mulheres x homens",survivalRatePerClassAndColumn(X, 1, 2) )
print("Classe 2 : mulheres x homens",survivalRatePerClassAndColumn(X, 2, 2) )
print("Classe 3 : mulheres x homens",survivalRatePerClassAndColumn(X, 3, 2) )
print("")
print("Tripulantes : mulheres x homens",totalPersonsPerClassAndColumn(X, 0, 2) )
print("Classe 1 : mulheres x homens",totalPersonsPerClassAndColumn(X, 1, 2) )
print("Classe 2 : mulheres x homens",totalPersonsPerClassAndColumn(X, 2, 2) )
print("Classe 3 : mulheres x homens",totalPersonsPerClassAndColumn(X, 3, 2) )


```

Output:
```
Tripulantes : 0.23955
Classe 1 : 0.62462
Classe 2 : 0.41404
Classe 3 : 0.25212

Tripulantes :criancas x adultos (None, 0.23955)
Classe 1 :criancas x adultos (1.0, 0.61755)
Classe 2 :criancas x adultos (1.0, 0.36015)
Classe 3 :criancas x adultos (0.34177, 0.24083)

Tripulantes : mulheres x homens (0.86957, 0.22274)
Classe 1 : mulheres x homens (0.97241, 0.34444)
Classe 2 : mulheres x homens (0.87736, 0.13966)
Classe 3 : mulheres x homens (0.45918, 0.17255)

Tripulantes : mulheres x homens (23, 862)
Classe 1 : mulheres x homens (145, 180)
Classe 2 : mulheres x homens (106, 179)
Classe 3 : mulheres x homens (196, 510)
```

Se olharmos para a taxa de sobrevivencia geral por classe, a classe que menos sobreviveu foi a classe de tripulantes, seguida pela classe 3. Por outro lado, se olharmos para a separaçao por sexo, temos que a taxa de sobrevivencia feminina da classe de tripulantes foi extremamente alta. A primeiro momento, isso parece incoerente. Mas se olharmos para o total de homens e mulheres da taxa de tripulantes, conseguimos ver que a quantidade de mulheres frente a homens na classe de tripulantes 'e significativamente menor, fazendo com que a alta taxa de sobrevivencia de mulheres dessa classe nao altere a taxa de sobrevivencia geral da classe. 

<!-- grade_id: cell33-a00, points: 1, grade: true, solution: true -->
